<?xml version="1.0" encoding="UTF-8"?>

<section xml:id="lesson2_4" xmlns:xi="http://www.w3.org/2001/XInclude">
  <title>Writing Quantum Programs with Qiskit</title>
  <p> So far, we have been exploring quantum circuits by hand, and using the circuit builder. Now,
    let's take a look at how we can build simple circuits using <term>Qiskit</term> (pronounced
    "kiss-kit"), a popular quantum programming library and toolkit from IBM. </p>

  <p>Let's start with the simplest circuit: a <c>NOT</c> gate.</p>

  <listing xml:id="not-gate">
    <title><c>NOT</c> Gate in Qiskit</title>
    <program xml:id="not-gate-program" interactive="activecode" language="python3">
      <code>
        from qiskit import QuantumCircuit

        circuit1 = QuantumCircuit(1)
        circuit1.x(0)
        print(circuit1.draw(fold=-1))
      </code>
    </program>
  </listing>

  <p>And that's it! Let's walk through this code line by line to understand it properly. We start by
    importing the <c>QuantumCircuit</c> class from the <c>qiskit</c> library: </p>

  <program language="python3">
    <code>
      from qiskit import QuantumCircuit
    </code>
  </program>

  <p>Next, we initialize a new circuit with 1 qubit by passing <c>1</c> as a parameter to the
    constructor, and call it <c>circuit1</c>:</p>
  <program language="python3">
    <code>
      circuit1 = QuantumCircuit(1)
    </code>
  </program>

  <p>Now, we add a <c>NOT</c> gate, using the <c>.x()</c> function. The <c>NOT</c> gate is also known as an <c>X</c>
    gate, which is why this is the name of the function. By passing the parameter <c>0</c> we specify
    that we want to add the <c>NOT</c> gate to the 0th qubit (qubits are 0-indexed!):</p>
  <program language="python3">
    <code>
      circuit1.x(0)
    </code>
  </program>

  <p>Finally, we draw the circuit:</p>
  <program language="python3">
    <code>
      print(circuit1.draw())
    </code>
  </program>

  <exercise label="qiskit-exercise-1">
    <title> Building Circuits with <c>NOT</c> in Qiskit </title>
    <idx>NOT circuit in Qiskit</idx>
    <statement>
      <p> Complete the code below to create a circuit with <em>two qubits</em> and a <c>NOT</c> gate on
        each. </p>
    </statement>
    <program interactive="activecode" language="python3">
      <code>
        from qiskit import QuantumCircuit

        def two_nots():
          # TODO: Pass in the correct parameter to create a circuit with 2 qubits
          circuit1 = QuantumCircuit(...)

          # TODO: Write code to add a NOT gate to both qubits

          # Draw the circuit
          print(circuit1.draw())
          return circuit1
      </code>
      <tests>
      from unittest.gui import TestCaseGui
      from qiskit import QuantumCircuit

      class myTests(TestCaseGui):
        def testOne(self):
          output = two_nots()

          self.assertEqual(type(output), QuantumCircuit, msg="Must return a QuantumCircuit")
          self.assertEqual(output.num_qubits, 2, msg="Circuit must have 2 qubits")
    </tests>
    </program>
  </exercise>

  <p> We can now include the 2-qubit gates we've learned, <c>SWAP</c> and <c>CNOT</c>:</p>

  <listing xml:id="swap-and-cnot-gates">
    <title><c>SWAP</c> and <c>CNOT</c> Gate in Qiskit</title>
    <program xml:id="swap-and-cnot-gates-program" interactive="activecode" language="python3">
      <code>
        from qiskit import QuantumCircuit

        circuit2 = QuantumCircuit(3)
        circuit2.x(0)
        circuit2.swap(0, 1)
        circuit2.cx(1, 2)

        print(circuit2.draw())
      </code>
    </program>
  </listing>

  <exercise label="qiskit-exercise-2">
    <title> Building Circuits with <c>NOT</c> in Qiskit </title>
    <idx>NOT circuit in Qiskit</idx>
    <statement>
      <p> Complete the code below to create a circuit with <em>two qubits</em> and a <c>NOT</c> gate on
        each. </p>
    </statement>
    <program interactive="activecode" language="python3">
      <code>
        from qiskit import QuantumCircuit

        def two_nots():
          # TODO: Pass in the correct parameter to create a circuit with 2 qubits
          circuit1 = QuantumCircuit(...)

          # TODO: Write code to add a NOT gate to both qubits

          # Draw the circuit
          print(circuit1.draw())
          return circuit1
      </code>
      <tests>
      from unittest.gui import TestCaseGui
      from qiskit import QuantumCircuit

      class myTests(TestCaseGui):
        def testOne(self):
          output = two_nots()

          self.assertEqual(type(output), QuantumCircuit, msg="Must return a QuantumCircuit")
          self.assertEqual(output.num_qubits, 2, msg="Circuit must have 2 qubits")
    </tests>
    </program>
  </exercise>
</section>
